```python
UPVD 2026
```

# Prise en main de l'environnement de travail

Suivez les
[instructions](http://nicolas.thiery.name/Enseignement/M1-ISD-AlgorithmiqueAvancee/ComputerLab/README.html#telechargement-et-depot-des-tps)
pour accéder à l'environnement de travail et téléchargez le TP.

Puis ouvrez cette première feuille Jupyter `00-PriseEnMain`.

## Prise en main de Jupyter

Cette feuille est dédiée à la prise en main des feuilles de travail Jupyter (que la
plupart d'entre vous connaisse), ainsi que du mécanisme de correction semi-automatique
que nous utiliserons dans ce cours.

Une analyse de données prend la forme d'un document narratif expliquant les objectifs,
hypothèses et étapes de l'analyse: quelles sont les données, qu'est-ce que l'on souhaite
calculer, pourquoi, comment, quels sont les résultats et quelles conclusions on en tire.
Ce document est typiquement écrit en anglais pour qu'il puisse être partagé avec le plus
grand nombre.

Depuis quelques années, les feuilles de travail (notebook) Jupyter sont un des moyens
prisés pour rédiger de telles analyses de données et mener les calculs sous-jacents. Il
s'agit en effet de documents permettant de mêler narration, interaction, calcul,
visualisation et programmation.

Le document que vous êtes en train de lire est une feuille Jupyter. On peut y mettre des
calculs comme ci-dessous. Pour exécuter le calcul, cliquez dans la *cellule* ci-dessous
et tapez Shift-Enter:

```python
1 + 1
```

Utilisez maintenant la cellule ci-dessous pour calculer 1+2; puis réutilisez la même
cellule pour calculer 3×4:

```python

```

Vous rencontrerez aussi des cellules à compléter comme la suivante dans lesquelles vous
remplacerez la ligne «# Remplacez cette ligne par votre code» par le calcul à faire.
Calculez cinq fois sept:

<!-- grade_id: cell-d8b61dce71965f3e, points: 0, grade: true, solution: true -->
```python
# ENTRER VOTRE CODE A LA PLACE DE CES 2 LIGNES
raise NotImplementedError()
```

<!-- grade_id: cell-b35d0d71e7e10dd3, points: 1, grade: true, solution: true -->
> VOTRE REPONSE A LA PLACE DE CE TEXTE (format texte ou markdown)

On peut ausi mettre des formules mathématiques dans une cellule:
$$\frac 1 {1-\frac 1z} = \sum_{i=0}^\infty \frac 1 {z^i}$$

Notez que certaines cellules de ce document comme celle-ci sont en lecture-seule. Vous ne
pouvez pas les modifier. En revanche, vous pouvez toujours insérer de nouvelles cellules.

Insérez une nouvelle cellule ci-dessous, et mettez y la formule $E=mc^2$:

Indications:

- Boutton + de la barre de menu de la feuille
- sélectionnez la cellule
- changez son type en `MarkDown` (voir la barre de menu de la feuille)
- double-cliquez sur cette cellule ou la précédente pour voir comment insérer des
  formules mathématiques en latex.

- Lancez la visite guidée de l'interface Jupyter.<br> Indication: Menu `Aide` ->
  `Visite de l'interface utilisateur`.
- Consultez les raccourcis claviers.<br> Indication: Menu `Aide` -> `Raccourcis clavier`.

## Correction semi-automatique avec nbgrader

Certaines feuilles de travail seront notées. Les notes de certaines séances ultérieures
contribueront à votre moyenne. Cette première séance est une séance d'entraînement; les
notes seront indicatives pour que vous preniez en main le mécanisme.

Une partie de la correction est manuelle: vos enseignants regarderont vos réponses et
attribueront des points à chacune d'entre elles. Le gros de la correction sera
automatique, s'appuyant sur des tests similaires à ceux du premier semestre.

Voici un exemple tout bête. Dans la cellule suivante, calculez la somme de `3` et de `4`,
et stockez le résultat dans la variable `s`:

<!-- grade_id: cell-d8b61dce71965f3g, points: 0, grade: false, solution: true -->
```python
# ENTRER VOTRE CODE A LA PLACE DE CES 2 LIGNES
raise NotImplementedError()
```

La commande suivante est un test qui vérifie votre réponse; vous reconnaîtrez le `assert`
(cette fois en minuscule) que nous utilisions en C++ au premier semestre:

<!-- grade_id: cell-d8b61dce71965f3h, points: 1, grade: true, solution: false -->
```python
assert s == 7
```

<!--Validez votre feuille en cliquant sur le bouton `Validate`; cela en
exécute une copie dans l'ordre en vérifiant tous les tests
automatiques.!-->

Suivez les
[instructions](http://nicolas.thiery.name/Enseignement/M1-ISD-AlgorithmiqueAvancee/ComputerLab/README.html#telechargement-et-depot-des-tps)
pour déposer votre travail sur GitLab et obtenir les résultats de la correction
automatique.

:::{important}

Les résultats de la correction automatique sont indicatifs. Mais ce que vous allez
développer servira dans tous les TPs suivants. Vous devez viser un zéro faute :-) En
sus, la qualité du code de ce TP et de sa documentation sera l'un des éléments de la
notation du TP 2.
:::

:::{admonition} Apprenez à utiliser Jupyter comme un pro!
:class: important

Il sera essentiel, pour ce cours et pour la suite, de devenir un utilisateur expert de
Jupyter, notamment en ce qui concerne l'utilisation des raccourcis. Pour cela, vous avez
par exemple à votre disposition ce
[tutoriel](https://tutoriels-jupyter.pages.in2p3.fr/lite/lab?path=tutoriel/index.ipynb),
à faire hors séance.
:::

## Au boulot!

Maintenant que vous avez appris les rudiments de Jupyter et de l'environnement de
travail; il est temps de passer au [TP](01-Introduction.md) lui-même!